// This shopping list program isn't compiling! 
// Use your knowledge of generics to fix it.

<<<<<<< Updated upstream
// I AM NOT DONE

||||||| Stash base
// Execute `rustlings hint generics1` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE

=======
// Execute `rustlings hint generics1` or use the `hint` watch subcommand for a hint.

>>>>>>> Stashed changes
fn main() {
    let mut shopping_list: Vec<&str> = Vec::new();
    shopping_list.push("milk");
}

